---
title: 'Types of Data'
description: 'Chapter description goes here.'
---

## Example coding exercise

```yaml
type: NormalExercise
key: e8c1edbe67
lang: python
xp: 100
skills: 2
```

This is an example exercise.

`@instructions`


`@hint`


`@pre_exercise_code`
```{python}

```

`@sample_code`
```{python}

```

`@solution`
```{python}

```

`@sct`
```{python}

```

---

## Variable Types

```yaml
type: VideoExercise
key: ab1a5203c7
xp: 50
```

`@projector_key`
11614ef99ebdf40b629637458e8ee6b2
